```python
import string
import re
import numpy as np
from numpy import dot
from numpy.linalg import norm
import hashlib
import math
import pandas as pd 
import matplotlib.pyplot as plt
from pylab import barh,plot,yticks,show,grid,xlabel,figure,cla,close
from nltk.tokenize import WordPunctTokenizer

from sklearn import datasets
from sklearn.metrics import accuracy_score
from sklearn.model_selection import train_test_split
from sklearn.feature_extraction.text import CountVectorizer, TfidfVectorizer
from sklearn.naive_bayes import MultinomialNB
from sklearn import svm
from sklearn.metrics import accuracy_score
from sklearn.feature_selection import SelectFromModel
from sklearn.feature_selection import SelectKBest, chi2

import spacy

df = pd.read_csv('datasets/twitter_rep_dem_data_small.csv')

handle_remotion = lambda doc: re.subn(r'@\w+','', doc.lower())[0]
df['Tweet'] = df['Tweet'].map(handle_remotion)

simple_tokenizer = lambda doc: " ".join(WordPunctTokenizer().tokenize(doc))
df['Tweet'] = df['Tweet'].map(simple_tokenizer)

train_data, test_data = train_test_split(df, test_size=0.3, random_state=42)

nlp = spacy.load('en_core_web_md')
```

## Q1.

You are studying Politics and want to know the type of language politicians use when they speak. To do this, you obviously went to twitter, since most politicians have become fans of making use of social media to target their electorate. You want to know the type of vocabulary politicians use, so you will start with the simplest approach you know and just extract "meaningful" n-grams with TF-IDF.


### Q1.a)

Implement a function to output the top k n-grams retrieved by the TF-IDF algorithm.

<!-- grade_id: cell-17155c56392d07fd, points: 0, grade: false, solution: true -->
```python
def get_k_best_tfidf(data, ngram_range, k):
    
    # vectorizer_truncated = ...
    # ...
    # feature_names = ...
    # return feature_names
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4279426c694d6d6e, points: 2, grade: true, solution: false -->
```python
assert(hashlib.sha256(get_k_best_tfidf(train_data.Tweet, (1,1), 100)[42].encode()).hexdigest() == "935f68319d4f227e02bfd54a0ddf85b8a242e42a4277aa5ef5eaab691710924e")
assert(hashlib.sha256(get_k_best_tfidf(train_data.Tweet, (1,2), 100)[42].encode()).hexdigest() == "866ab23f02991a070ef7b06fc39ab77a0e9e840db9d0dabb204735b109d4a52c")
```

```python
print(get_k_best_tfidf(train_data.Tweet, (1,4), 100)[0:10])
```

### Q1.b)

You got your results and were not entirely satisfied. You notice that some of your selected features are somewhat similar and that, no matter the party of a given politician, they all like to mutter the same things, hoping to get your vote and then do what they promised to the people who financed their campaigns. Fortunately, you remember learning something about stopwords.

Re-implement the previous method but avoiding choosing stopwords.

<!-- grade_id: cell-0dcac4c40253e789, points: 0, grade: false, solution: true -->
```python
def get_k_best_tfidf_plus(data, ngram_range, k, language=None):
    
    # vectorizer_truncated = ...
    # ...
    # feature_names = ...
    # return feature_names
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

Verify that the results hold as before when you don't provide a language

<!-- grade_id: cell-9b55532e713fb5f1, points: 1, grade: true, solution: false -->
```python
assert(hashlib.sha256(get_k_best_tfidf_plus(train_data.Tweet, (1,1), 100)[42].encode()).hexdigest() == "935f68319d4f227e02bfd54a0ddf85b8a242e42a4277aa5ef5eaab691710924e")
assert(hashlib.sha256(get_k_best_tfidf_plus(train_data.Tweet, (1,2), 100)[42].encode()).hexdigest() == "866ab23f02991a070ef7b06fc39ab77a0e9e840db9d0dabb204735b109d4a52c")
```

```python
print(get_k_best_tfidf_plus(train_data.Tweet, (1,4), 100)[0:10])
```

Now let's tests some of our outputted words:

<!-- grade_id: cell-1c81e595725f5b5c, points: 1, grade: true, solution: false -->
```python
assert(hashlib.sha256(get_k_best_tfidf_plus(train_data.Tweet, (1,1), 100, language='english')[42].encode()).hexdigest() == "b2c66dfd45dd07f58d2246b695660b220ee429c6a8ba95f27e4c57f02e9d6ac8")
assert(hashlib.sha256(get_k_best_tfidf_plus(train_data.Tweet, (1,2), 100, language='english')[42].encode()).hexdigest() == "b2c66dfd45dd07f58d2246b695660b220ee429c6a8ba95f27e4c57f02e9d6ac8")
```

```python
print(get_k_best_tfidf_plus(train_data.Tweet, (1,4), 100, language='english')[0:10])
```

## Q2.

After taking a careful look, you decide that, although your results are good, due to your ambitious nature, you still want to improve and get more meaningful features.

Repeat Q1, but this time, instead of truncating the TF-IDF features, use the chi-squared method to get the relevant features, given the political stance of the authors of each tweet. Implement a function that outputs the most important features according to this approach.

<!-- grade_id: cell-ca3cb6e3d721fc9e, points: 0, grade: false, solution: true -->
```python
def get_k_best_chi(data, labels, ngram_range, k):
    
    # vectorizer = ...
    # ...
    # X_train = ...
    # feature_names = ...
    # ch2 = ...
    # stat_X_train_chi = ...
    # most_important = ...
    # return most_important
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-dd51064dc35465dc, points: 2, grade: true, solution: false -->
```python
assert(hashlib.sha256(get_k_best_chi(train_data.Tweet, train_data.Party, (1,1), 100)[42].encode()).hexdigest() == "d6e21286621a8586f7e54720ab5a39c93acc2f4b8fba7a16ec1c24d69a08c613")
assert(hashlib.sha256(get_k_best_chi(train_data.Tweet, train_data.Party, (1,2), 100)[42].encode()).hexdigest() == "c23b31a0179b550f8a18fb06bc52a26e32333540369073d646da5c84a4dc341f")
```

```python
most_important_features = get_k_best_chi(train_data.Tweet, train_data.Party, (1,2), 100)[0:10]
print(most_important_features)
```

## Q3.

You now feel confident enough to verify  if these features are in fact meaningful for classifying your data, so you decide to plot them.

```python
for feature in most_important_features:
    print('Documents that contains the word(s) "%s"' % feature)
    print('----')
    docs = train_data.Tweet.str.lower().str.contains(feature)
    print(str(train_data.Party[docs].value_counts()) + '\n\n')
```

Which of these statements is true?

**A** - If the word "2020census" appears in our text, we know that it must be a Democrat Tweet 

**B** - Republican tweets contains much more specific and therefore relevant vocabulary

**C** - Chi-squared is implemented to ignore stopwords

**D** - If we see the word "betterway" the author is probably Republican 

<!-- grade_id: cell-5e448e3e9a1a11fb, points: 0, grade: false, solution: true -->
```python
#answer = # Answer with a one character string, capitalized

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3c7d72311b6a3941, points: 2, grade: true, solution: false -->
```python
expected_hash = '3f39d5c348e5b79d06e842c114e6cc571583bbf44e4b0ebfda1a01ec05745d43'
assert hashlib.sha256(str(answer).encode()).hexdigest() == expected_hash
```

## Q4.

Back when you were taking your degree you had some algebra classes where you learned about some pretty cool concepts. You have also heard about the possibility of using them to tackle some practical issues in data science, and you wonder if they would be useful to what you are doing now. You now want try methods like SVD and PCA, but, because you don't want to waste a long time waiting for your experiments to run, you're going to start out with using a smaller dataset - the iris dataset from scikit-learn.

```python
%matplotlib inline

iris = datasets.load_iris()
X = iris.data[:, :2]  # we only take the first two features.
y = iris.target

# import some data to play with
iris = datasets.load_iris()
X = iris.data[:, :2]  # we only take the first two features.
y = iris.target

x_min, x_max = X[:, 0].min() - .5, X[:, 0].max() + .5
y_min, y_max = X[:, 1].min() - .5, X[:, 1].max() + .5

plt.figure(2, figsize=(6,4))
plt.clf()
# Plot the training points
plt.scatter(X[:, 0], X[:, 1], c=y, cmap=plt.cm.Set1, edgecolor='k')
plt.xlabel('Sepal length')
plt.ylabel('Sepal width')
```

You first verified that the data is not that easy to separate taking two of your features, so you now want to see the results for the first few components. Implement SVD decomposition to return a reduced version of your data and see the results you get.

<!-- grade_id: cell-615fa1a7ad60dda0, points: 0, grade: false, solution: true -->
```python
def svd_decomposition(A, n_components):

    # AT = ...
    # AAT = ...
    # ATA = ...
    
    # lmbU, U = ...
    # lmbV, V = ...

    #S_shape = ...
    #lmb = ...
    #S = ...
    
    #reduction = ...
    #return reduction
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
X_reduced = svd_decomposition(iris.data, 2)

x_min, x_max = X_reduced[:, 0].min() - .5, X_reduced[:, 0].max() + .5
y_min, y_max = X_reduced[:, 1].min() - .5, X_reduced[:, 1].max() + .5

plt.figure(2, figsize=(6,4))
plt.clf()

# Plot the training points
plt.scatter(X_reduced[:, 0], X_reduced[:, 1], c=y, cmap=plt.cm.Set1, edgecolor='k')
plt.xlabel('SVD 1')
plt.ylabel('SVD 2')
```

You should see a much better separation with SVD than you saw before. If you don't, something is not right. Let's check if the components you obtained are right:

<!-- grade_id: cell-baf6a0f82935e3ac, points: 2, grade: true, solution: false -->
```python
X_reduced = svd_decomposition(iris.data, 2)

assert math.isclose(X_reduced[:, 0].sum(axis=0), 1151.87905039)
assert math.isclose(X_reduced[:, 1].sum(axis=0), 38.48015192)
```

## Q5.

After testing SVD you also want to test the other cool-sounding method, PCA. Implement it and see how its results compare with the the ones obtained for SVD.

<!-- grade_id: cell-f882c24a78354337, points: 0, grade: false, solution: true -->
```python
def pca_reduction(A, n_components):
    
    # u = ...
    # covA = ...
    # lmbV, V = ...
    # V_red = ...
    # reduction = ...
    # return reduction  

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
X_reduced_pca = pca_reduction(iris.data - iris.data.mean(axis=0), 2)

x_min, x_max = X_reduced_pca[:, 0].min() - .5, X_reduced_pca[:, 0].max() + .5
y_min, y_max = X_reduced_pca[:, 1].min() - .5, X_reduced_pca[:, 1].max() + .5

plt.figure(2, figsize=(6,4))
plt.clf()
# Plot the training points
plt.scatter(X_reduced_pca[:, 0], X_reduced_pca[:, 1], c=y, cmap=plt.cm.Set1, edgecolor='k')
plt.xlabel('PC 1')
plt.ylabel('PC 2')
```

And finally, to make sure everything is correct:

<!-- grade_id: cell-3fabdbdf0da61899, points: 2, grade: true, solution: false -->
```python
assert math.isclose(X_reduced_pca[42, 0], -2.998296442832354, abs_tol=1e-6)
assert math.isclose(X_reduced_pca[42, 1], 0.3343075745907761, abs_tol=1e-6)
```

## Q6.

Now that you have mastered the art of using algebra methods for dimensionality reduction, pick the only true statement in the following questions:


### Q6.a)

**E** - SVD and PCA are essentialy the same method

**F** - SVD requires us to center the data before computing a decomposition

**G** - PCA can be solved by applying SVD internally

**H** - PCA and SVD are not the same method but return the same principal components

<!-- grade_id: cell-dec9bff37d676bcb, points: 0, grade: false, solution: true -->
```python
#answer = # Answer with a one character string, capitalized

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0abe1a2e6fb7cbe0, points: 2, grade: true, solution: false -->
```python
expected_hash = '333e0a1e27815d0ceee55c473fe3dc93d56c63e3bee2b3b4aee8eed6d70191a3'
assert hashlib.sha256(str(answer).encode()).hexdigest() == expected_hash
```

### Q6.b)

**I** - The sum of the eigenvalues of a data matrix equals that data variance

**J** - SVD requires us to center the data before computing a decomposition

**K** - The sum of the eigenvalues of the covariance matrix of a given data matrix equals that data variance

**L** - The singular value decomposition generates four matrices that multiplied reconstruct the original data.

<!-- grade_id: cell-f5662a10917b237a, points: 0, grade: false, solution: true -->
```python
#answer = # Answer with a one character string, capitalized

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-33058c9dbf1e5c97, points: 2, grade: true, solution: false -->
```python
expected_hash = '86be9a55762d316a3026c2836d044f5fc76e34da10e1b45feee5f18be7edb177'
assert hashlib.sha256(str(answer).encode()).hexdigest() == expected_hash
```

## Q7.

After a long day looking at politicians tweeting as often as CMTV finds a new disaster to report and then trying to separate flowers with the powers of mathematics you wonder what is the meaning of life after all. You mistakenly say it aloud and a guy suspiciously looking like Tomas Mikolov starts muttering a long sequence of numbers. The thought that maybe numbers and vectors are really the answer goes through your head. Then it comes to you that you have heard in the past about something called embeddings and how they are able to incorporate the meaning of words into vectors and you want to test it.

Before doing so, you need a way to measure the distance between these vectors and therefore, between words. Start by implementing a function that is able to measure the difference of cosine similarity between the vectors of two words.

```python
def vec(s):
    return nlp.vocab[s].vector
```

<!-- grade_id: cell-24db69517366bdc8, points: 0, grade: false, solution: true -->
```python
def cosine_diff(vector_pair_1, vector_pair_2):
    
    #v1, v2 = vector_pair_1
    #if ...:
    #    cosine_1 = ...
    #else:
    #    cosine_1 = ...
    #v3, v4 = vector_pair_2
    #if ...:
    #    cosine_2 = ...
    #else:
    #    cosine_2 = ...
        
    #return cosine_1 - cosine_2
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f52da173a887ccbd, points: 2, grade: true, solution: false -->
```python
assert(math.isclose(cosine_diff((vec('man'), vec('boy')), (vec('man'), vec('girl'))), 0.13031119, abs_tol=1e-7))
```

## Q8.

What the heck? It seems to be true. These weird looking vectors are indeed able to say that the word boy is closer to the word man than the word girl is. Hmmm.. you decide to investigate a little bit further. How do you know they're not foooling you? Let's see if they are able to infer some very simple associations between words. Implement the necessary code to do so.

```python
def cosine(v1, v2):
    if norm(v1) > 0 and norm(v2) > 0:
        return dot(v1, v2) / (norm(v1) * norm(v2))
    else:
        return 0.0

def spacy_closest(token_list, vec_to_check, n=10, dont_include_list=[]):
    return sorted([x for x in token_list if x not in dont_include_list],
                  key=lambda x: cosine(vec_to_check, vec(x)),
                  reverse=True)[:n]
```

```python
vectorizer = TfidfVectorizer()
vectorizer.fit_transform(train_data.Tweet)

tweet_vocab = vectorizer.vocabulary_
```

<!-- grade_id: cell-8df0ab8b60111598, points: 0, grade: false, solution: true -->
```python
def word_vector_arithmetic(token_list, positive_vecs=[], negative_vecs=[], n=10, dont_include_list=[]):
 
    #vec_to_check = ...
    #for vec in positive_vecs:
    #    if vec_to_check is None:
    #        vec_to_check = ...
    #    else:
    #        vec_to_check += ...
    #for vec in negative_vecs:
    #    vec_to_check -= ...
    #return spacy_closest(...)

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-efb61c0c575def1a, points: 2, grade: true, solution: false -->
```python
closest_word = word_vector_arithmetic(tweet_vocab.keys(),
                                      positive_vecs=[vec('prince'), vec('woman')],
                                      negative_vecs=[vec('man')],
                                      dont_include_list=['prince', 'woman', 'man'])[0]

assert(closest_word == 'princess')
```

Maybe that Tomas Mikolov fellow wasn't muttering those numbers in vain.. Well, but to know that you will have to keep studying about the topic :)